# functions3.py


# We tried to write a function format_name, that takes a firstname and a 
# lastname and returns a fullname, but we made a mistake.
# Can you fix it?


### I AM NOT DONE


def format_name(firstname: str, lastname: str) -> str:
    fullname = firstname + " " + lastname

format_name("Jane", "Doe")

print("The name is", fullname)
