# functions2.py


# You now know how to create a function.
# We now want to create a function that turns hours into minutes.
# The function should take an int (hours) and return another int (minutes).
# Please finish the function.
# Note, the type declarations are optional.


### I AM NOT DONE


def hours_to_minutes(hours: int) -> int:
    return

# Don't modify code below

def test_hours_to_minutes():
    assert hours_to_minutes(1) == 60
    assert hours_to_minutes(5) == 300
    assert hours_to_minutes(5) == 360

