# variables3.py


# There is nothing technically wrong with the variable below,
# But we have assigned the wrong value
# Can you fix it WITHOUT changing the assignment on line 11?


### I AM NOT DONE

number = "5"     # don't change this line


# Don't change code below

def test_number_is_5():
    assert number == 5