# variables3.py

# We have defined a variable number with the value of "5"
# Now we want to add 5 to it, but it is not possible to add a number 
# to a string
# Can you redefine the variable number, so the code works?

### I AM NOT DONE

number = "5"     # don't change this line

print("our number is " + number)

number =           # change this line to fix the code

number_plus_5 = number + 5  # don't change this line

# Don't change code below

def test_number_is_5():
    assert number == 5