# classes1.py

# We are trying to create a class that models a person.
# It is supposed to take and save a persons firstname, lastname and their age.
# Please finish writing the Person class and create two person objects.
# Uncomment the def __str__ and see how the print statements change.


### I AM NOT DONE


class Person:
    def __init__(self, firstname: str, lastname: str, age: int) -> None:
        self.firstname = firstname
        # finish implementation here
        
    # def __str__(self) -> str:
    #     return f"{self.lastname}, {self.firstname}({self.age})"



person1 = 

person2 = 

print(person1)
print(person2)
