# conditions3.py


# Write two functions, return_smaller and return_bigger
# return_smaller is supposed to compare two numbers, a and b, and return the smaller one.
# return_bigger is supposed to compare two numbers, a and b, and return the bigger one.


### I AM NOT DONE

def return_smaller(a: int, b: int) -> int:
    return

def return_bigger(a: int, b: int) -> int:
    return


# Don't modify code below

def test_return_smaller():
    assert return_smaller(1, 2) == 1
    assert return_smaller(5, 2) == 2
    assert return_smaller(2, 2) == 2

def test_return_bigger():
    assert return_bigger(1, 2) == 2
    assert return_bigger(5, 2) == 5
    assert return_bigger(2, 2) == 2

