# conditions2.py


# We are printing a custom message, depending on what day of the week we have.
# Can you finish the code?


### I AM NOT DONE

weekday = 

if weekday == "Monday":
    print("Oh no. The weekend is over")
elif weekday == "Tuesday":
    print("Tuesday, Bluesday")
elif weekday == "Wednesday":
    
else:
    print("That's not one of the days of the week :O")