# lists1.py

# We've tried to create a list called ingredients below, but there is an error
# Can you fix it for us?
# Feel free to look at the code that is printing the ingredients below
# and see if you can understand it

### I AM NOT DONE

ingredients = ["beans", "teriyaki sauce", "apricot jam", "plum jam", "coconut milk",

# Don't modify code below, 
# but you can take a look and try to understand it

for ingredient in ingredients:
    print(ingredient)