# lists3.py

# We want to add to number_list in different ways.
# Look up how to add to lists in python and follow
# the instructions in the comments below.
# Don't reassign number_list
# Hint: list methods: append, insert, extend


### I AM NOT DONE


number_list = [0, 1, 2, 4, 5]

# Add the number 6 at the end of number_list
number_list

# Add the number 3 inbetween number 2 and 4 of number_list
number_list

# Add contents of higher_numbers at the end of number_list
higher_numbers = [7, 8, 9]
number_list

# Don't modify code below 

if number_list != [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]:
    raise ValueError("number_list should be [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]")