# dictionaries1.py

# We created a dictionary that stores the number of wins of Alex and Sam.
# Finish writing the if statement, so that winner is assigned to the person with 
# more wins. 


### I AM NOT DONE


score = {
    "Alex": 3,
    "Sam": 2
}

if score > score:
    winner =
else:
    winner =

# Don't modify code below

if winner != "Alex":
    raise ValueError("But Alex won...")

