# strings4.py


# We want to replace the word orange in the order defined below with banana.
# For that we want to use the banana variable, also defined below.
# Only modify the line where you defined modified_order.
# Be sure to use the order and banana variables.
# Hint: Look at string methods: strip and replace


### I AM NOT DONE

banana = "  banana  "       # don't modify

order = "I want three apples and one orange"    # don't modify

modified_order = 

# Don't modify code below

assert modified_order == "I want three apples and one banana"