# for_loops3.py

# We are writing the while loop from while_loops2 with a for loop instead
# Reminder: We're trying to count till 13, but skip number 7
# Can you finish the code?
# Hint: try out the continue statement, and the range function

### I AM NOT DONE



import time


for count in range:
    if count == 7:
        raise ValueError("continue here")

    print(count)

    time.sleep(1)
