# basic_types3.py



# Finish the code below, to try out more basic calculation.
# Check out the variable names, to see what the variable is supposed to store.

### I AM NOT DONE


x = 3

x_plus_three = 

x_minus_five = 

x_times_two = 

x_divided_by_two = 

x_to_power_of_three = 


# Don't modify code below
# After you finish, you can take a look at it, 
# if you want to get a preview of if statements.

if(x_plus_three != x+3):
    raise ValueError(f"{x_plus_three} is not {x} plus three.")
elif(x_minus_five != x-5):
    raise ValueError(f"{x_minus_five} is not {x} minus five.")
elif(x_times_two != x*2):
    raise ValueError(f"{x_times_two} is not {x} times two.")
elif(x_divided_by_two != x/2):
    raise ValueError(f"{x_divided_by_two} is not {x} divided by two.")
elif(x_to_power_of_three != x**3):
    raise ValueError(f"{x_to_power_of_three} is not {x} to the power of three.")
else:
    print(f"{x_plus_three} is {x} plus three.")
    print(f"{x_minus_five} is {x} minus five.")
    print(f"{x_times_two} is {x} times two.")
    print(f"{x_divided_by_two} is {x} divided by two.")
    print(f"{x_to_power_of_three} is {x} to the power of three.")
