# basic_types1.py

# Below, we tried to assign three strings to three variables, but the code won't run.
# Can you fix it?


### I AM NOT DONE

my_string = "I am a string

my_second_string = 'I'm also a string'

my_third_string = "I said, "I am a string""


print(my_string)
print(my_second_string)
print(my_third_string)