# basic_types4.py



# Follow the instructions below, to fix the assertion errors


### I AM NOT DONE

# use float() or int(), so that these assertions works
int_5 = 5.0
float_5 = 5

assert isinstance(int_5, int)      
assert isinstance(float_5, float)

# use round() to make this assertion pass.
assert 6.67 == 20/3

# Fix the assertions, by replacing the None with the correct result.
# Ask yourself: What do these operators do? How could they be useful?
assert 3//2 == None
assert 3%2 == None