# while_loops1.py


# We've written a countdown for new years for you.
# Try it out.
# What happens if you add line 19, that is currently commented out?

### I AM NOT DONE

import time

counter = 

while counter > 0:
  print(counter)

  time.sleep(1)
  
  # counter = counter - 1

print("Happy New Year!!")