# while_loops1.py


# We've written a countdown for new years for you.
# Try it out.
# What happens if you add line 19, that is currently commented out?

### I AM NOT DONE

import time

counter = 

while counter > 0:
  print(counter)

  time.sleep(0.5)     # when you're done, remove this line 
  
  counter = counter - 1 # ask yourself why this line is here

print("Happy New Year!!")