# while_loops2.py


# We're trying to count till a number of your choosing,
# but stop the loop after that number.
# Can you do that WITHOUT changing the conditional of the while loop?
# Hint: try out the break statement


### I AM NOT DONE


import time


counter = 1

while True:         # DON'T CHANGE THIS LINE
    print(counter)

    time.sleep(0.5)     # when you're done, remove this line 
  
    counter = counter + 1
